---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook
### Spring Term 2022-23
#### Prepared by Dr Haider Raza

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function to pass a list of numbers and calculate the exponential of all elements in the list format. 
`Hint`: take input as list and inside the function covert it into numpy array then perform the function. `Example` : input `[1, 4, 8]` and and output should be:  `[0.0, 2.0, 3.0]` 

<!-- grade_id: myLog_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myExp(mylist):
    out_list = [np.exp(elem) for elem in mylist]
    return out_list
    raise NotImplementedError()
```

```python
in_array = [1, 2, 4, 6]
myExp(in_array)
```

Output:
```
[2.718281828459045, 7.38905609893065, 54.598150033144236, 403.4287934927351]
```

<!-- grade_id: myLog_test, points: 1, grade: true, solution: false -->
```python
"""run these test cases"""
# Test case 1
my = [1, 2, 4, 6]
assert myExp(my) == [2.718281828459045, 7.38905609893065, 54.598150033144236, 403.4287934927351]

# Test case 2
my = [1, 10, 100]
assert myExp(my) == [2.718281828459045, 22026.465794806718, 2.6881171418161356e+43]
```

---

## Question 2:

Write a function to find the k number of the smallest values from a NumPy array and remove duplicate when returning as a list. 


For example: 
`nSmallest([23, 12, 1, 3, 1, 4, 0, 4, 5, 6],3) = [0, 1, 3]`

<!-- grade_id: myfibonacci_ans, points: 0, grade: false, solution: true -->
```python
def nSmallest(array, k):
    out = []
    array = list(set(array))
    for i in range(k):
      a = np.min(array)
      array.remove(a)
      out.append(a)
    return out
    raise NotImplementedError()
```

```python
nSmallest([10, 10, 2, 5, 4, 5, 6, 8, 2], 2)
```

Output:
```
[2, 4]
```

<!-- grade_id: myfibonacci_test, points: 1, grade: true, solution: false -->
```python
"""run these test cases"""
assert nSmallest([23, 12, 1, 3, 1, 4, 0, 4, 5, 6],3) == [0, 1, 3]
assert nSmallest([100, 200, 2000, 30, 110],4) == [30, 100, 110, 200]
assert nSmallest([10, 10, 2, 5, 4, 5, 6, 8, 2], 2) == [2, 4]

```

---
## Question 3: 

 Write a program, which accept a list with `nan`s in it and count the maximum number of consecutive repeated `nan`.
 Example: In the `a = [nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16]` and answer is 2.

<!-- grade_id: max_nan_ans, points: 0, grade: false, solution: true -->
```python
from numpy import nan
from math import isnan
def max_rep_nan(in_list):
    counter = 0
    max = 0
    for elem in in_list:
      if isnan(elem):
        counter +=1
        if max<counter:
          max = counter
      else:
        counter = 0
    return max 
    raise NotImplementedError()
```

```python
a = np.array([nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16])
max_rep_nan(a)
```

Output:
```
2
```

<!-- grade_id: max_rep_test, points: 0, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""
a = [nan, nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 3

a = [nan, nan, nan, nan, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 5

a = [nan, 10, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 1

a = []
assert max_rep_nan(a) == 0
```

## Pandas Data Frame 


```python
from sklearn import datasets
import pandas as pd
```

```python
# import some data to play with
diabetes = datasets.load_diabetes()
X = pd.DataFrame(diabetes.data[:, :])  # we are taking all features.
y = diabetes.target
print(X)
print(y)
```

Output:
```
0         1         2         3         4         5         6  \
0    0.038076  0.050680  0.061696  0.021872 -0.044223 -0.034821 -0.043401   
1   -0.001882 -0.044642 -0.051474 -0.026328 -0.008449 -0.019163  0.074412   
2    0.085299  0.050680  0.044451 -0.005671 -0.045599 -0.034194 -0.032356   
3   -0.089063 -0.044642 -0.011595 -0.036656  0.012191  0.024991 -0.036038   
4    0.005383 -0.044642 -0.036385  0.021872  0.003935  0.015596  0.008142   
..        ...       ...       ...       ...       ...       ...       ...   
437  0.041708  0.050680  0.019662  0.059744 -0.005697 -0.002566 -0.028674   
438 -0.005515  0.050680 -0.015906 -0.067642  0.049341  0.079165 -0.028674   
439  0.041708  0.050680 -0.015906  0.017282 -0.037344 -0.013840 -0.024993   
440 -0.045472 -0.044642  0.039062  0.001215  0.016318  0.015283 -0.028674   
441 -0.045472 -0.044642 -0.073030 -0.081414  0.083740  0.027809  0.173816   

            7         8         9  
0   -0.002592  0.019908 -0.017646  
1   -0.039493 -
```

---
## Question 4:

Write a function to find the max number in the each column in the diabetes dataset and return the output as a list. Please make sure to round the output to two decimal points. 

<!-- grade_id: filter_1_ans, points: 0, grade: false, solution: true -->
```python
def my_max_min(X):
    """Write a function to find the max number in the each column and return 
    as a list."""
    # YOUR CODE HERE
    return [round(elem, 2) for elem in X.max().to_list()]
    raise NotImplementedError()
```

<!-- grade_id: filter_1_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
assert my_max_min(X) == [0.11, 0.05, 0.17, 0.13, 0.15, 0.2, 0.18, 0.19, 0.13, 0.14]
```

---
## Question 5:


Write a function to calculate the `mean` and `standard deviation` of diabetes dataset (first column only). `Hint`: You can use NumPy aggregate functions. Create a list and append value to it and return the output. 

<!-- grade_id: stats_ans, points: 0, grade: false, solution: true -->
```python
def stats(X):
    """Write a functon to calculate mean and standard deviation of IRIS dataset (last column only)."""
    return [np.mean(X.iloc[:,0]), np.std(X.iloc[:,0])]
    raise NotImplementedError()
```

```python
stats(X)
```

Output:
```
[-3.6396225400041895e-16, 0.04756514941545118]
```

```python

```

<!-- grade_id: stats_test, points: 2, grade: true, solution: false -->
```python
"""Check your output"""
# Test case
assert stats(X) == [-3.6396225400041895e-16, 0.04761904761904939]
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-16-fcd712ddb311> in <module>
      1 """Check your output"""
      2 # Test case
----> 3 assert stats(X) == [-3.6396225400041895e-16, 0.04761904761904939]

AssertionError:
```

```python

```