---
# Welcome to CE880
### This is your week-3 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

We have already loaded `iris` dataset for you. Now, write a function to calculate the `median` on the 1st column of the dataset. 
`Hint`: To find the median, you first need to sort the values in your sample. But, before that we have sliced, the first column of `iris` data into a variable name `sample` and data type is a list. Then pass `sample` list as a parameter in you `my_median` function.

The `median` of a sample of numeric data is the value that lies in the middle when we sort the data. The data may be sorted in ascending or descending order, the median remains the same. Please also consider what to do if there is no middle value.

Please don't use `np.median` function. Write your own custom function

<!-- grade_id: my_median_ans, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets
iris = datasets.load_iris()
sample = list(iris.data[:,0])

def my_median(sample):
    """Function for calculating median"""
    sample.sort()
    median = sample[int(round(len(sample)/2, 0))]
    return median
    raise NotImplementedError()
```

<!-- grade_id: my_median_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
import math
assert math.isclose(my_median(sample), 5.8, rel_tol=0.05)
```

---

## Question 2:

Write a function to Caluate Interquartile Range (IQR) on the second column of `iris` dataset, which you have loaded in the question 1.
`hint`: We have created a new variable `sample2` and sliced the second column of the `iris` data. Please explore `np.percentile` function

<!-- grade_id: my_quartile_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
sample2 = list(iris.data[:,1])
def my_quartile(sample2):
    """Caluate Interquartile Range of the given list."""
    iqr = np.percentile(sample2,q=75) - np.percentile(sample2, q=25)
    return iqr

    raise NotImplementedError()
```

<!-- grade_id: my_quartile_test, points: 2, grade: true, solution: false -->
```python
"""Check that correct answers here"""
import math
assert math.isclose(my_quartile(sample2), 0.5, rel_tol=0.05)
```

---
## Question 3: 

Write a function to calculate the `standard deviation` on the third column of `iris` dataset. We have created a new variable `sample3` and sliced the third column of the iris data. Please create a function that first calculate the variance then calculate `standard deviation`. Please do not use inbuilt `numpy.var` and `numpy.std` functions. Write your own codes. 

<!-- grade_id: my_std_ans, points: 0, grade: false, solution: true -->
```python
import math
sample3 = list(iris.data[:,2])
def my_std(sample3, ddof=0):
    """Write a program to calculate the standard deviation of 3rd column of iris dataset."""
    
    # YOUR CODE HERE
    var = 0
    mean = sum(sample3)/len(sample3)
    for elem in sample3:
      var += (elem-mean)**2
    var = var/len(sample3)
    return math.sqrt(var)
    raise NotImplementedError()
```

<!-- grade_id: my_std_test, points: 1, grade: true, solution: false -->
```python
"""Test you solution"""
assert math.isclose(my_std(sample3), 1.759, rel_tol=0.05)
```

---
## Question 4:

Write a program to find the mean of the 4th column of `iris` dataset. 

<!-- grade_id: my_mean_ans, points: 0, grade: false, solution: true -->
```python
sample4 = list(iris.data[:,3])

def my_mean(sample4):
    """write a function to find the mean of the 4th column of iris data"""
    return sum(sample4)/len(sample4)
    raise NotImplementedError()
```

<!-- grade_id: my_mean_test, points: 1, grade: true, solution: false -->
```python
"""Test you solution"""
assert math.isclose(my_mean(sample4), 1.19, rel_tol=0.05)
```

## Question 5

Write a Pandas program to calculate the correlation between column 1 and column 3 of `iris` dataset. `Hint` use `from scipy import stats` to find a Pearson correlation coefficient. 

<!-- grade_id: iris_correlation_ans, points: 0, grade: false, solution: true -->
```python
import itertools
from scipy import stats
import pandas as pd

sample1 = list(iris.data[:,0])
sample2 = list(iris.data[:,2])
def iris_correlation(sample1, sample2):
    """calculate the correlation"""
    # YOUR CODE HERE
    return stats.pearsonr(sample1, sample2)[0]
    raise NotImplementedError()
```

```python
"""Test you solution"""
assert math.isclose(iris_correlation(sample1, sample2), 0.87, rel_tol=0.05)
```

```python

```

```python

```